<a href="https://colab.research.google.com/github/Kn23982/kn23982-CE880-Lab-coursework/blob/main/week2/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "kn23982"
COLLABORATORS = ""
```

---

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook
### Autumn Term 2023-24
#### Prepared by Dr Haider Raza

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/).


---
## Question 1:

Write a Python function that takes a list of numbers and returns a list of their cosine values.

By solving this question correctly you will get 1 point.

<!-- grade_id: cosine_ans, points: 0, grade: false, solution: true -->
```python
import math
def calculate_cosine_of_list(numbers):
    """Write a Python function that takes a list of numbers and returns a list of their cosine values """
    result = []
    for num in numbers:
      cosine_result = round(math.cos(num),2)
      result.append(cosine_result)
      return result
    raise NotImplementedError()
```

```python
numbers = [0, -2, 44, -61, -30, 44]
cosine_result = calculate_cosine_of_list(numbers)
print(cosine_result)
```

Output:
```
[1.0]
```

<!-- grade_id: cosine_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""

# Test case 1
my = [0, math.pi / 2, math.pi, 100, 200, 230]
assert calculate_cosine_of_list(my) == [1.0, 0.0, -1.0, 0.86, 0.49, -0.79]

# Test case 2
my = [0, 2, 4, 6, -30, 400]
assert calculate_cosine_of_list(my) == [1.0, -0.42, -0.65, 0.96, 0.15, -0.53]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-9-88234f19fc4d> in <cell line: 5>()
      3 # Test case 1
      4 my = [0, math.pi / 2, math.pi, 100, 200, 230]
----> 5 assert calculate_cosine_of_list(my) == [1.0, 0.0, -1.0, 0.86, 0.49, -0.79]
      6 
      7 # Test case 2

AssertionError:
```

---

## Question 2:

Given an array of integers, find two numbers such that they add up to a specific target sum. Return the indices of the two numbers. For example, given the array `[2, 7, 11, 15]` and the target `9`, the function should return `[0, 1]` since `2 + 7 = 9`.

By solving this question correctly you will get 1 point.

<!-- grade_id: twosum_ans, points: 0, grade: false, solution: true -->
```python
def two_sum(nums, target):
    """Given an array of integers, find two numbers such that they add
    up to a specific target sum. Return the indices of the two numbers."""
    for i in range (len(nums)):
        if nums[i] > target:
             continue
        for j in range(i + 1, len(nums)):
             if nums[i] + nums[j] == target:
               return [i,j]
    raise NotImplementedError()
```

```python
# Example for students
nums = [2, 7, 11, 15]
target = 9
result = two_sum(nums, target)
print(result)
```

Output:
```
[0, 1]
```

<!-- grade_id: twosum_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""

# Case 1
nums = [17, 12, 31, 6]
target = 37
assert two_sum(nums, target) == [2, 3]

# Case 2
nums = [1, 2, 4, 6, 30, 40]
target = 31
assert two_sum(nums, target) == [0, 4]
```

---
## Question 3:

Given an array containing all integers from `0` to `n` except for one missing number, find the missing number. For example, given `[3, 0, 1]`, the missing number is `2`.

By solving this question correctly you will get 1 point.

<!-- grade_id: missnum_ans, points: 0, grade: false, solution: true -->
```python

def find_missing_number(nums):
    """Given an array containing all integers from 0 to n except for one missing number,
    find the missing number. For example, given [3, 0, 1],
    the missing number is 2."""
for num in nums:
      n = len(nums)
      excepted_sum = n*(n+1)//2
      actual_sum = sum(nums)
      missing_number = excepted_sum - actual_sum
  return missing_number
raise NotImplementedError()
```

Output:
```
File "<tokenize>", line 10
    return missing_number
    ^
IndentationError: unindent does not match any outer indentation level
```

```python
nums = [3, 0, 1]
result = find_missing_number(nums)
print(result)
```

<!-- grade_id: missnum_test, points: 1, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
# Case 1
nums = [3, 0, 1]
assert find_missing_number(nums) == 2

# Case 2
nums = [4, 5, 6, 2, 0, 1]
assert find_missing_number(nums) == 3

```

## Pandas Data Frame


```python
from sklearn import datasets
import pandas as pd
# Create a sample DataFrame
data = {
    'Name': ['Alice', 'Bob', 'Charlie', 'David'],
    'Age': [25, 30, 28, 35],
    'City': ['New York', 'Los Angeles', 'Chicago', 'Houston']
}
df = pd.DataFrame(data)
df
```

---
## Question 4:

Write a function to Find the age of the person named `Bob` and `Bob` is not in the Data Frame then return `0`.

By solving this question correctly you will get 2 points.

<!-- grade_id: findele_ans, points: 0, grade: false, solution: true -->
```python
def find_element(dataframe, name):
    """Write a function to Find the age of the person named "Bob" """

    if name in dataframe['Name'].values:
      bob_age = dataframe.loc[dataframe['Name'] == name, 'Age'].values[0]
      return bob_age
    else:
      return 0

    raise NotImplementedError()
```

<!-- grade_id: findele_test, points: 2, grade: true, solution: false -->
```python
assert find_element(df, 'Bob') == 30
assert find_element(df, 'Charlie') == 28
assert find_element(df, 'CE880') == 0
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-70-6523ef7b75f9> in <cell line: 1>()
----> 1 assert find_element(df, 'Bob') == 30
      2 assert find_element(df, 'Charlie') == 28
      3 assert find_element(df, 'CE880') == 0

NameError: name 'df' is not defined
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of `Age`. `Hint` Create a list and append value to it and return the output.

By solving this question correctly you will get 1 point.

```python

```

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of Age from the
    above dataframe df and round the values to 2 decimals points ."""

    age_values = x['Age'].values.tolist()
    mean_age = round(np.mean(age_values),2)
    std_age = round(np.std(age_values),2)
    return mean_age , std_age
    raise NotImplementedError()
```

<!-- grade_id: stats_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
x = stats(df)
y= [29.5, 4.2]
np.testing.assert_allclose(x, y, rtol=0.05, atol=0)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-72-714e70460d8d> in <cell line: 3>()
      1 """Check your output"""
      2 # Test case
----> 3 x = stats(df)
      4 y= [29.5, 4.2]
      5 np.testing.assert_allclose(x, y, rtol=0.05, atol=0)

NameError: name 'df' is not defined
```